---
title: 'Importing Data'
description: 'Import data from the UN Comtrade database.'
---

## Import Data

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Import the ___un_comtrade_data.csv_

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}
read_csv(file = "un_comtrade_data.csv"
```

`@sct`
```{r}

```
